<a href="https://colab.research.google.com/github/ioiein/Network-science/blob/main/assignment_knowledge_graph_embedding/assignment.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Assignment — Knowledge Graph Embedding

In this assignment we will see how to use the [TorchKGE](https://github.com/torchkge-team/torchkge) library for building knowledge graphs and its embeddings.

```python
import numpy as np
import pandas as pd
pd.options.display.max_colwidth = 100
from zlib import adler32
import requests
```

### Task 1. Dataset exploration (1 points)

To begin with we are going to need a knowledge graph, so let us load a standard knowledge graph dataset called _Freebase-15k-237_.

```python
df = pd.read_csv(
    "https://raw.githubusercontent.com/netspractice/ml-on-graphs/main/datasets/freebase-237-merged-and-remapped.csv", 
    names=['h', 'r', 't'])
```

```python
df = df[~df.h.str.startswith('/') & ~df.t.str.startswith('/')]
df[::1000].head(10)
```

Output:
```
h  ...                            t
0     queens college, city university of new york  ...                 carol leifer
1103                                 sharon stone  ...                 jacob vargas
2210                                        japan  ...                        tokyo
3329                                     syphilis  ...                         male
4447                        michael carlisle hall  ...                  mena suvari
5555                          alice in wonderland  ...             english language
6653                                        nauru  ...             english language
7743                                   ben harper  ...                     marriage
8840              the perks of being a wallflower  ...     united states of america
9960                               vice president  ...  scripps howard broadcasting

[10 rows x 3 columns]
```

There is h — head (also subject), r — relation (also predicat, label), t — tail (also object). The shape of the dataset is

```python
df.shape
```

Output:
```
(282560, 3)
```

Let us check the number of unique entities and unique relations.

Write a funtion `n_ent_rel` that takes a dataset and returns a number of unique entities and unique relations.


<!-- grade_id: cell-a1c37b647b2bd8d4, points: 0, grade: false, solution: true -->
```python
def n_ent_rel(df):
    # YOUR CODE HERE
    # raise NotImplementedError()
    return df["h"].append(df["t"]).nunique(), df["r"].nunique()
```

<!-- grade_id: cell-e297130eb1b840de, points: 0.25, grade: true, solution: false -->
```python
n_ent, n_rel = n_ent_rel(df)
assert adler32(str(n_ent + n_ent).encode()) == 50266365
n_ent, n_rel
```

Output:
```
(13506, 228)
```

We can look at all facts about any entity via pandas API

```python
df[df.h ==  'aleksandr pushkin'].head()
```

Output:
```
h  ...                           t
4107   aleksandr pushkin  ...                        poet
14992  aleksandr pushkin  ...  johann wolfgang von goethe
19525  aleksandr pushkin  ...              russian empire
40237  aleksandr pushkin  ...                      moscow
47411  aleksandr pushkin  ...         william shakespeare

[5 rows x 3 columns]
```

```python
df[(df.h ==  'aleksandr pushkin') & (df.t == 'nikolai gogol')]
```

Output:
```
h  ...              t
145219  aleksandr pushkin  ...  nikolai gogol

[1 rows x 3 columns]
```

```python
df[(df.h ==  'nikolai gogol') & (df.t == 'aleksandr pushkin')]
```

Output:
```
h  ...                  t
86362   nikolai gogol  ...  aleksandr pushkin
269512  nikolai gogol  ...  aleksandr pushkin

[2 rows x 3 columns]
```

Let us try to find some facts in this dataset. For example, what is Harrison Ford's nationality? ("harrison ford" in the dataset)

Write a function `harrison_ford_nationality` that takes a dataset and returns the nationality.

_Hint: use `pandas.Series.str.contains` method_

<!-- grade_id: cell-bf2eb40c10e9971d, points: 0, grade: false, solution: true -->
```python
def harrison_ford_nationality(df):
    # YOUR CODE HERE
    # raise NotImplementedError()
    return df[df["h"].str.contains("harrison ford") & df["r"].str.contains("nationality")]["t"].values[0]
```

<!-- grade_id: cell-e20681a97039cdf6, points: 0.25, grade: true, solution: false -->
```python
assert adler32(harrison_ford_nationality(df).encode()) == 1948191013
harrison_ford_nationality(df)
```

Output:
```
'united states of america'
```

More tricky question: who are film directors of movies where Harrison Ford was?

Write a function `made_films_with_harrison_ford` that returns a set of directors' names.

<!-- grade_id: cell-7ea1e48135c759d5, points: 0, grade: false, solution: true -->
```python
def made_films_with_harrison_ford(df):
    # YOUR CODE HERE
    # raise NotImplementedError()
    s1 = set(df.loc[(df["h"] == "harrison ford")&(df["r"].str.contains("film")), "t"])
    s2 = set(df.loc[(df["t"].isin(s1))&(df["r"].str.contains("director")), "h"])
    return s2
```

<!-- grade_id: cell-f3dbc03094dcbb5b, points: 0.5, grade: true, solution: false -->
```python
directors = made_films_with_harrison_ford(df)
assert adler32(str(sorted(directors)).encode()) == 3798415057
directors
```

Output:
```
{'francis ford coppola',
 'george lucas',
 'jon favreau filmography',
 'kathryn bigelow',
 'mike nichols',
 'ridley scott',
 'robert zemeckis',
 'steven spielberg'}
```

### Task 2. TransE embedding model (2 points)

TransE (translations in the embedding space) is a method which models relationships by interpreting them as translations operating on the low-dimensional embeddings of the entities.

Let us implement the model using TorchKGE — a Python module for knowledge graph embedding relying solely on Pytorch.

```python
!pip install torchkge==0.16.25 -q
```

Output:
```
[?25l[K     |██████▉                         | 10kB 14.0MB/s eta 0:00:01[K     |█████████████▋                  | 20kB 9.5MB/s eta 0:00:01[K     |████████████████████▍           | 30kB 4.9MB/s eta 0:00:01[K     |███████████████████████████▏    | 40kB 4.4MB/s eta 0:00:01[K     |████████████████████████████████| 51kB 1.6MB/s 
[?25h
```

```python
import torch
from torch import nn
import torch.nn.functional as F
from torchkge import KnowledgeGraph
from torchkge.utils import Trainer
from torchkge.evaluation import LinkPredictionEvaluator
from torchkge.models.interfaces import TranslationModel
from torch.nn.init import xavier_uniform_
```

<img src='https://github.com/netspractice/ml-on-graphs/raw/main/assignment_knowledge_graph_embedding/kge.png' width=300>

The basic idea behind the model is that the functional relation induced by the $r$-labeled edges corresponds to a translation of the embeddings, i.e. we want that $h + r \approx t$ when ($h$, $r$, $t$) holds ($t$ should be a nearest neighbor of $h + r$), while $h + r$ should be far away from $t$ otherwise. Thereby, TransE scoring function is negative distance between $h+r$ and $t$

$$f(h, r, t) = -||h + r - t||$$

where $||\cdot||$ is L1 or L2 norm. The goal of the optimization procedure is learning optimal embeddings, such that the scoring function is able to assign high scores to positive statements and low scores to statements unlikely to be true. The most common loss function is Margin Loss that can be represented as

$$L = \text{ReLU}\left(\gamma - f(h, r, t) + f(h', r, t')\right)$$

where $\gamma > 0$ is a "margin" hyperparameter, the triplet $(h, r, t)$ is in a training set (positive) and the $(h', r, t')$ is in a set of *corrupted* (negative) triplets, that are composed of training triplets with either the head or tail replaced by a random entity (but not both at the same time).

We will use TorchKGE TranslationModel as a base for TransE class, here are some auxiliary methods: 

* `normalize_parameters` applies `L2` normalization for training stability
* `get_embeddings` returns the tensors representing entities and relations
* `lp_prep_cands` gets entities and relations embeddings for link prediction training process
* `forward` returns scores for positive triplets (`h`, `r`, `t`) and negative triplets (`nh`, `r`, `nt`)

```python
class BaseTransE(TranslationModel):
    def __init__(self, num_entities, num_relations, dim=100):
        super(BaseTransE, self).__init__(num_entities, num_relations, dissimilarity_type='L2')
        self.num_entities = num_entities
        self.num_relations = num_relations
        self.dim = dim
        self.ent_embeddings = nn.Embedding(num_entities, self.dim)
        xavier_uniform_(self.ent_embeddings.weight.data)
        self.rel_embeddings = nn.Embedding(num_relations, self.dim)
        xavier_uniform_(self.rel_embeddings.weight.data)
        
        self.normalize_parameters()
        self.rel_embeddings.weight.data = F.normalize(self.rel_embeddings.weight.data, p=2, dim=1)

    def normalize_parameters(self):
        self.ent_embeddings.weight.data = F.normalize(self.ent_embeddings.weight.data, p=2, dim=1)

    def get_embeddings(self):
        self.normalize_parameters()
        return self.ent_embeddings.weight.data, self.rel_embeddings.weight.data
        
    def lp_prep_cands(self, h_idx, t_idx, r_idx):
        b_size = h_idx.shape[0]

        h_emb = self.ent_embeddings(h_idx)
        t_emb = self.ent_embeddings(t_idx)
        r_emb = self.rel_embeddings(r_idx)

        candidates = self.ent_embeddings.weight.data.view(1, self.num_entities, self.dim)
        candidates = candidates.expand(b_size, self.num_entities, self.dim)

        return h_emb, t_emb, candidates, r_emb
    
    def forward(self, h, t, nh, nt, r):
        return self.scoring_function(h, t, r), self.scoring_function(nh, nt, r)
```

Write a function `scoring_function` that takes a tensor `h`, `t`, `r` with triplets of the shape (number of triplets) and returns its L2 scores. For training stability, apply L2 normalization (`F.normalize`  with `p=2`) on entities embeddings before computing the scores. Note that `h`, `t`, `r` are tensors with integer IDs of entities and relations.

<!-- grade_id: cell-10c450a856899a4f, points: 0, grade: false, solution: true -->
```python
class TransE(BaseTransE):
    def scoring_function(self, h, t, r):
        # YOUR CODE HERE
        # raise NotImplementedError()
        h_ent_embd = F.normalize(self.ent_embeddings(h), p=2, dim=-1)
        t_ent_embd = F.normalize(self.ent_embeddings(t), p=2, dim=-1)
        r_ent_embd = self.rel_embeddings(r)
        return -torch.linalg.norm(h_ent_embd + r_ent_embd - t_ent_embd, ord=2, dim=-1)
```

<!-- grade_id: cell-89f5ccbf84380079, points: 0, grade: true, solution: false -->
```python
torch.manual_seed(0)
model = TransE(num_entities=5, num_relations=2, dim=2)
model.ent_embeddings.weight.data = torch.randn(5, 2)
model.rel_embeddings.weight.data = torch.randn(2, 2)

h = torch.LongTensor([1, 1, 1])
r = torch.LongTensor([0, 1, 1])
t = torch.LongTensor([2, 2, 3])

with torch.no_grad():
    scores = model.scoring_function(h, t, r).numpy()
assert scores.shape == (3, ) # 3 triplets
assert adler32(str(scores.round(2)).encode()) == 628491186
```

Complete the class `MarginLoss` so that computes the losses among scores and returns an average loss. Positive and negative scores are tensors of the shape (number of triplets).

<!-- grade_id: cell-3fb7cb717cb420cb, points: 0, grade: false, solution: true -->
```python
class MarginLoss(nn.Module):
    def __init__(self, margin):
        super().__init__()
        self.margin = margin
    def forward(self, positive_scores, negative_scores):
        # YOUR CODE HERE
        # raise NotImplementedError()
        return torch.mean(F.relu(self.margin - positive_scores + negative_scores))
```

<!-- grade_id: cell-b28b48d0354bf7be, points: 0, grade: true, solution: false -->
```python
loss_fn = MarginLoss(margin=0.5)
torch.manual_seed(0)
positive_scores = torch.randn(8)
negative_scores = torch.randn(8)
loss = loss_fn(positive_scores, negative_scores)
assert loss.shape == torch.Size([])
assert round(loss.item(), 2) == 0.79
```

Let us look at the WikiDataSet that presents country-specific subgraphs of Wikidata.

```python
url = 'https://raw.githubusercontent.com/netspractice/ml-on-graphs/main/datasets/countries_edges.tsv'
open('countries_edges.tsv', 'wb').write(requests.get(url).content)
url = 'https://raw.githubusercontent.com/netspractice/ml-on-graphs/main/datasets/countries_entities.tsv'
open('countries_entities.tsv', 'wb').write(requests.get(url).content)
url = 'https://raw.githubusercontent.com/netspractice/ml-on-graphs/main/datasets/countries_relations.tsv'
open('countries_relations.tsv', 'wb').write(requests.get(url).content);
```

```python
edges = pd.read_csv('countries_edges.tsv', sep='	').values
entity_labels = pd.read_csv('countries_entities.tsv', sep='	', index_col=0).label.values
relation_labels = pd.read_csv('countries_relations.tsv', sep='	', index_col=0).label.values
```

```python
edges_labeled = np.stack([entity_labels[edges[:, 0]], 
                          entity_labels[edges[:, 1]], 
                          relation_labels[edges[:, 2]]], axis=1)

df = pd.DataFrame(edges_labeled, columns=['h', 't', 'r'])[['h', 'r', 't']]
df.head()
```

Output:
```
h                    r              t
0  Belgium              country        Belgium
1  Belgium  diplomatic relation     Luxembourg
2  Belgium  diplomatic relation          Italy
3  Belgium  diplomatic relation        Germany
4  Belgium  diplomatic relation  Liechtenstein
```

Here are unique relations with the number of triplets

```python
df.r.groupby(df.r).count()
```

Output:
```
r
applies to jurisdiction                                        1
basic form of government                                       3
biological process                                             1
capital                                                       34
capital of                                                    25
contains administrative territorial entity                    94
country                                                     1339
country of citizenship                                         5
country of origin                                              1
culture                                                        3
designated as terrorist by                                     9
different from                                               113
diplomatic relation                                         6145
enclave within                                                 3
facet of                                                      13
family                 
```

Convert our dataset into a TorchKGE knowledge graph

```python
kg = KnowledgeGraph(
    pd.DataFrame(edges_labeled, columns=['from', 'to', 'rel']))
```

Split the dataset into train and test set. What differs from the standard method of randomly sampling N points to make up our test set, is that our data points are two entities linked by some relationship, and we need to take care to ensure that all entities are represented in train and test sets by at least one triple.



```python
kg_train, kg_test = kg.split_kg()
```

Create required objects: model, loss, optimizer and trainer.

```python
model = TransE(kg_train.n_ent, kg_train.n_rel, dim=64)
criterion = MarginLoss(margin=0.5)
optimizer = torch.optim.Adam(model.parameters(), lr=0.001)

trainer = Trainer(
    model, criterion, kg_train, n_epochs=200, 
    batch_size=2048, optimizer=optimizer)
```

During training, positive triplets are selected, negative triplets are generated, margin loss is calculated and then the gradient step is performed.

```python
trainer.run()
```

Output:
```
HBox(children=(FloatProgress(value=0.0, max=200.0), HTML(value='')))
```

Let us evaluate our model on the link prediction task

<!-- grade_id: cell-e6b0a4e916e8c503, points: 2, grade: true, solution: false -->
```python
evaluator = LinkPredictionEvaluator(model, kg_test)
evaluator.evaluate(b_size=128)

assert evaluator.hit_at_k(k=10)[0] > 0.3
assert evaluator.mean_rank()[0] < 70
assert evaluator.mrr()[0] > 0.1
```

Output:
```
HBox(children=(FloatProgress(value=0.0, description='Link prediction evaluation', max=17.0, style=ProgressStyl…
```

```python
print('Hit@10: {:.4f}'.format(evaluator.hit_at_k(k=10)[0]))
print('Mean Rank: {:.4f}'.format(evaluator.mean_rank()[0]))
print('MRR: {:.4f}'.format(evaluator.mrr()[0]))
```

Output:
```
Hit@10: 0.3729
Mean Rank: 65.1693
MRR: 0.1476
```

`Hit@k` indicates how many times in average a true triple was ranked in the top-k.

`Mean Rank` is a mean rank of the true entity when replacing alternatively head and tail in any fact of the dataset.

`MRR` is an average of mean recovery rank for head and tail replacement.

### Task 3. 2D entity embedding (2 points)

```python
from sklearn.decomposition import TruncatedSVD
from sklearn.cluster import k_means
import matplotlib.pyplot as plt
```

Similarly to node embedding visualization, we can plot 2D embeddings of entities of knowledge graph.

Write a function `ent_embeddgins` that takes the model, transforms entity embeddings into 2-dimenstional space via SVD, then finds 2 clusters via k-means and finally returns a tuple: np.array with embeddings and 0/1 labels of clusters.

<!-- grade_id: cell-c79123779965d790, points: 0, grade: false, solution: true -->
```python
def ent_embeddgins(model):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c4d5a66741d71c62, points: 2, grade: true, solution: false -->
```python
ent_emb, labels = ent_embeddgins(model)
assert ent_emb.shape == (1810, 2)
assert labels.shape == (1810,)
n_label1 = (labels == 1).sum()
n_label0 = (labels == 0).sum()
assert 0.14 < min(n_label1, n_label0) / max(n_label1, n_label0) < 0.16
```

```python
plt.figure(figsize=(8, 5))
colors = np.array(plt.cm.tab10.colors)
plt.scatter(ent_emb[:, 0], ent_emb[:, 1], s=10, c=colors[labels])
plt.title('Entity embeddings')
plt.show()
```

Let us look at some randomly chosen entities from each cluster. Here we use `kg.ent2ix` dictionary to get labels of entities.

```python
ent_labels = np.array(list(kg.ent2ix.keys()))
print('First cluster: ', ', '.join(
    np.random.permutation(ent_labels[labels == 0])[:10]))
print()
print('Second cluster: ', ', '.join(
    np.random.permutation(ent_labels[labels == 1])[:10]))
```

### Task 4. Nearest neighbors of an entity (2 points)

```python
from sklearn.neighbors import NearestNeighbors
```

Let us find a nearest neighbors of Belgium using embedding space.

Write a function `similar_countries` that takes a name of country, graph and model and returns a list with names of nearest countries. Use `model.get_embeddings()`.

<!-- grade_id: cell-fe878bfa8dc0c817, points: 0, grade: false, solution: true -->
```python
def similar_countries(name, kg, model):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0b8b02590a965cc3, points: 2, grade: true, solution: false -->
```python
similar = similar_countries('Belgium', kg, model)
assert 'Netherlands' in similar
similar
```

### Task 5. Relation prediction  (1.5 points)

Let us consider these facts:
1. Belgium shares border with France
2. Belgium shares border with Switzerland
3. Belgium shares border with Nigeria

Only the fact 1 is truth, but it seems that the fact 2 is more likely than 3.

```python
df[(df.h == 'Belgium') & (df.t == 'France')]
```

```python
df[(df.h == 'Belgium') & (df.t == 'Switzerland')]
```

```python
df[(df.h == 'Belgium') & (df.t == 'Nigeria')]
```

Now we can compare these facts using scoring function of TransE model to check our prior knowledge.

Write a function `belgium_facts` that takes a model, a graph and returns 3 values of scoring function for each fact. Use `model.scoring_function` with `torch.no_grad()`.

<!-- grade_id: cell-5feee5e866b961d0, points: 0, grade: false, solution: true -->
```python
def belgium_facts(model, kg):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a6b6e4e88ba6569f, points: 1.5, grade: true, solution: false -->
```python
scores = belgium_facts(model, kg)
assert scores[0] > scores[1] > scores[2]
scores
```

### Task 6. Tail prediction (1.5 points)

```python
from sklearn.neighbors import NearestNeighbors
```

Recall that the main idea of TransE is to learn $h + r \approx t$, so let us find the closest tails for the head + relation pair.

Write a function `soviet_tail` that takes a model, graph and finds 5 nearest neighbors for the pair $h$ is "Soviet Union" and $r$ is "founded by".

<!-- grade_id: cell-dbc103d80a938c68, points: 0, grade: false, solution: true -->
```python
def soviet_tail(model, kg):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ba34f9e5a695e8ca, points: 1.5, grade: true, solution: false -->
```python
tails = soviet_tail(model, kg)
assert 'Russian Soviet Federative Socialist Republic' in tails
tails
```

```python

```